<a href="https://colab.research.google.com/github/sbbzplt/QiskitFallFest_IZTECH/blob/master/Question_3_Bonus_Q.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Instructions

The places where you have enter code, to answer the questions, are marked with `# YOUR CODE HERE`. Once you have written your code you should remove the `raise NotImplementedError()` statement.

## Question 3 - Bonus

Suppose we have 4 guests among which there are enemies and We want to list all possible ways so that no enemies sit next to each other.

To represent who sits in each seat, we will use 2 qubits, hence 8 qubits in total. The four guests are labeled as 00, 01, 10, 11. 

As input, you will have the list of enemies. For instance, if input is ["00","01"], then 0 and 1 shouldn’t sit next to each other.

For instance `11 00 01 00` is not a valid arrangement.

We want to use Grover’s Search to list all possible seating arrangements. So, a marked state is the one in which each person has a seat and enemies do not sit next to each other.

You will write an oracle function to mark valid seating arrangements:

- You need to check if each guest appears exactly once
- You need to make sure that enemies do not sit next to each other
- Your oracle should output 1 in qubit numbered 7 to mark the valid states


You oracle should have the following properties:

* Inputs: List of enemy pairs
* Returns: Quantum circuit



<!-- grade_id: cell-7ce9d86f1511a847, points: 0, grade: false, solution: true -->
```python
from qiskit import *

def oracle(enemies):
    qc = QuantumCircuit(20,9)
    ### BEGIN SOLUTION
    
    # Your code here
    
    # Your code ends here
    ### END SOLUTION
    #uncompute
    
    return qc
    
```

```python
#This is used in the function below

def check_valid(counts,enemies):
    valid = True
    enemy_check = True
    
    for state in counts:
        message = ""
        guests = ["00","01","10","11"]
        occ = {guest: 0 for guest in guests}
        for i in range(4):
            for guest in guests:
                if state[2*i]==guest[0] and state[2*i+1]==guest[1]:
                    occ[guest]+=1
                    message += f"Guest {guest} sits in poisition {i}\n"      
        for i in range(3):
            if  state[2*i]==enemies[0][0] and state[2*i+1]==enemies[0][1] and state[2*i+2]==enemies[1][0] and state[2*i+3]==enemies[1][1]:
                enemy_check = False
                if state[8] == "1":
                    print(message)
                    print("Enemies sit next to each other at position ", i," and ", i+1,"\n" )
                    valid = False
            elif  state[2*i]==enemies[1][0] and state[2*i+1]==enemies[1][1] and state[2*i+2]==enemies[0][0] and state[2*i+3]==enemies[0][1]:
                enemy_check = False
                if state[8] == "1":
                    print(message)
                    print("Enemies sit next to each other at position ", i," and ", i+1,"\n" )
                    valid = False
        
        if enemy_check and list(occ.values()) == [1,1,1,1] and state[8]=="0":
            print(message)
            print("Arrangement is valid but you do not mark it\n")
            valid = False
        elif list(occ.values())!= [1,1,1,1] and state[8]=="1":
            print(message)
            print("Arrangement is not valid but you mark it\n")
            valid = False
    return valid
```

```python
# You can use this code to test your function

qc = QuantumCircuit(20,9)

#Create equal superposition state
for i in range(8):
    qc.h(i)
    

enemies = ["00","01"]
qc += oracle(enemies)

#We are measuring in reverse order
for i in range(9):
    qc.measure(i,8-i)
    
job = execute(qc,Aer.get_backend('qasm_simulator'),shots=1000)
counts = job.result().get_counts(qc)   
if(check_valid(counts,["00","01"])):
    print("Oracle is working")

```

Output:
```
Oracle is working
```

```python
#You can use this cell to count the number of operations

tp = transpile(qc, basis_gates = ['cx', 'u'])
print(tp.count_ops())
```

Output:
```
OrderedDict([('u', 3894), ('cx', 3692), ('barrier', 24), ('measure', 9)])
```

<!-- grade_id: cell-ee6caa56cfe85f1f, points: 1, grade: true, solution: false -->
```python
# hidden tests will be used for grading.
# If this cell results in an error, your implementation is incorrect
### BEGIN HIDDEN TESTS
# You can use this code to test your function

def test_oracle(enemies):
    qc = QuantumCircuit(20,9)

    #Create equal superposition state
    for i in range(8):
        qc.h(i)

    qc += oracle(enemies)

    #We are measuring in reverse order
    for i in range(9):
        qc.measure(i,8-i)

    job = execute(qc,Aer.get_backend('qasm_simulator'),shots=1000)
    counts = job.result().get_counts(qc)   

    assert check_valid(counts,enemies) == True

enemy_list = [["00","01"],["00","11"],["10","01"]]

for enemies in enemy_list:
    test_oracle(enemies)

tp = transpile(qc, basis_gates = ['cx', 'u'])
print(tp.count_ops())
### END HIDDEN TESTS

```

Output:
```
OrderedDict([('u', 3894), ('cx', 3692), ('barrier', 24), ('measure', 9)])
```