<a href="https://colab.research.google.com/github/sbbzplt/QiskitFallFest_IZTECH/blob/master/Question_3_Final_Q.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Instructions

The places where you have enter code, to answer the questions, are marked with `# YOUR CODE HERE`. Once you have written your code you should remove the `raise NotImplementedError()` statement.

## Question 3

Suppose we have 4 guests and we want to list all possible ways they can seat using a quantum circuit.

To represent who sits in each seat, we will use 2 qubits, hence 8 qubits in total. The four guests are labeled as 00, 01, 10, 11.

Example: 

State `10 00 01 11` implies that 
(10 (qubits 0,1), 00 (qubits 1,2), 01 (qubits 3,4), 11 (qubits 5,6))

- Guest 2 sits at position 0
- Guest 0 sits at position 1 
- Guest 1 sits at position 2
- Guest 3 sits at position 3

We want to use Grover’s Search to list all possible seating arrangements. So, a marked state is the one in which each person has a seat.

For instance `10 00 01 00` is not a valid arrangement since 11 is not seated and 00 is seated in two places.

You will write an oracle function to mark valid seating arrangements:

- You need to check if each guest appears exactly once
- Your oracle should output 1 in qubit numbered 7 to mark the valid states

The circuit has 20 qubits but you should try to use
- as few gates as possible
- as few qubits as possible

You oracle should have the following properties:

* Inputs: No input
* Returns: Quantum circuit



<!-- grade_id: cell-7ce9d86f1511a847, points: 0, grade: false, solution: true -->
```python
from qiskit import *

def oracle():
    qc = QuantumCircuit(20,9)
    ### BEGIN SOLUTION
    # Your code here

    

    # Your code ends here
    ### END SOLUTION    
    return qc
    
```

```python
#This is used in function below

def check_valid(counts):
    valid = True
    
    for state in counts:
        message = ""
        guests = ["00","01","10","11"]
        occ = {guest: 0 for guest in guests}
        for i in range(4):
            for guest in guests:
                if state[2*i]==guest[0] and state[2*i+1]==guest[1]:
                    occ[guest]+=1
                    message += f"Guest {guest} sits in poisition {i}\n"      
        if list(occ.values()) == [1,1,1,1] and state[8]=="0":
            print(message)
            print("Arrangement is valid but you do not mark it\n")
            valid = False
        elif list(occ.values())!= [1,1,1,1] and state[8]=="1":
            print(message)
            print("Arrangement is not valid but you mark it\n")
            valid = False
    return valid
```

```python
# You can use this code to test your function

qc = QuantumCircuit(20,9)

#Create equal superposition state
for i in range(8):
    qc.h(i)
#Call oracle
qc += oracle()

#We are measuring in reverse order
for i in range(9):
    qc.measure(i,8-i)
    
job = execute(qc,Aer.get_backend('qasm_simulator'),shots=1000)
counts = job.result().get_counts(qc)   
if(check_valid(counts)):
    print("Oracle is working")

```

Output:
```
Oracle is working
```

```python
#You can use this cell to count the number of operations 
from qiskit.converters import circuit_to_dag, dag_to_circuit
tp = transpile(qc, basis_gates = ['cx', 'u'])
print(tp.count_ops())
```

Output:
```
OrderedDict([('u', 313), ('cx', 228), ('barrier', 10), ('measure', 9)])
```

<!-- grade_id: cell-ee6caa56cfe85f1f, points: 1, grade: true, solution: false -->
```python
# hidden tests will be used for grading.
# If this cell results in an error, your implementation is incorrect
### BEGIN HIDDEN TESTS

qc = QuantumCircuit(20,9)

#Create equal superposition state
for i in range(8):
    qc.h(i)

qc += oracle()

#We are measuring in reverse order
for i in range(9):
    qc.measure(i,8-i)

job = execute(qc,Aer.get_backend('qasm_simulator'),shots=1000)
counts = job.result().get_counts(qc)   

assert check_valid(counts) == True
tp = transpile(qc, basis_gates = ['cx', 'u'])
print(tp.count_ops())
### END HIDDEN TESTS

```

Output:
```
OrderedDict([('u', 313), ('cx', 228), ('barrier', 10), ('measure', 9)])
```