<a href="https://colab.research.google.com/github/sbbzplt/QiskitFallFest_IZTECH/blob/master/Question_2_Final_Q.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Instructions

The places where you have enter code, to answer the questions, are marked with `# YOUR CODE HERE`. Once you have written your code you should remove the `raise NotImplementedError()` statement.

## Question 2

You are given a quantum circuit where the qubits numbered 0,1,2,3 are used for storing the input and qubits numbered 4 and 5 are used for the output.

Your task is to count the total number of "1"s in the input in modulo 3 and store it as a binary number in the output qubits.

Example:

- If the input state is 1010, the output should be 10 (qubit no 5 is 0 and qubit no 6 is 1)
- If the input state is 1011, the output should be 00 (qubit no 5 is 0 and qubit no 6 is 0)
- If the input state is 1000, the output should be 01 (qubit no 5 is 1 and qubit no 6 is 0)
- If the input state is 1111, the output should be 01 (qubit no 5 is 1 and qubit no 6 is 0)

The circuit has 20 qubits but you should try to use
- as few qubits as possible
- as few gates as possible

Note: The additional qubits you use do not have to be in 0 state at the end of the circuit.

* Inputs: A quantum circuit with 20 qubits
* Returns: Quantum circuit

<!-- grade_id: cell-0f4849fc0faf37f8, points: 0, grade: false, solution: true -->
```python
from qiskit import *

def count1(qc):
     ### BEGIN SOLUTION

    #Your code here

    #Your code ends here.
    ### END SOLUTION
    qc.measure(4,0)
    qc.measure(5,1)
   
    return qc
    
```

```python
# You can use this code to test your function

qc = QuantumCircuit(20,2)

#initialize to 1111
qc.x(0)
qc.x(1)
qc.x(2)
qc.x(3)

#Let's count occurence of "1"s in mod 3, Output should be 01
qc = count1(qc)

job = execute(qc,Aer.get_backend('qasm_simulator'),shots=1000)
counts = job.result().get_counts(qc)   

print(counts)

```

Output:
```
{'01': 1000}
```

```python
#You can use this cell to count the number of operations and number of qubits used

from qiskit.converters import circuit_to_dag, dag_to_circuit

def find_active_qubits(circuit):
    dag = circuit_to_dag(circuit)
    return 20-len([w for w in dag.idle_wires()])

tp = transpile(qc, basis_gates = ['cx', 'u'])
print(tp.count_ops())
print("Qubits used", find_active_qubits(qc))
```

Output:
```
OrderedDict([('u', 43), ('cx', 32), ('measure', 2)])
Qubits used 8
```

<!-- grade_id: cell-6ff4d14141f4705d, points: 1, grade: true, solution: false -->
```python
# hidden tests will be used for grading.
# If this cell results in an error, your implementation is incorrect
### BEGIN HIDDEN TESTS
def initialize(qc,inp):
    for i in range(len(inp)):
        if inp[i]=="1":
            qc.x(i)
    return qc
        
def generate_binary_strings(bit_count):
    binary_strings = []
    def genbin(n, bs=''):
        if len(bs) == n:
            binary_strings.append(bs)
        else:
            genbin(n, bs + '0')
            genbin(n, bs + '1')
    genbin(bit_count)
    return binary_strings

def count_occ(inp):
    c=0
    for i in inp:
        c+=int(i)
    return format(c%3, '02b')

from qiskit.converters import circuit_to_dag, dag_to_circuit

def find_active_qubits(circuit):

    dag = circuit_to_dag(circuit)
    return 20-len([w for w in dag.idle_wires()])

inputs = generate_binary_strings(4)
for inp in inputs:
    qc = QuantumCircuit(20,2)
    initialize(qc,inp)
    qc = count1(qc)
    result = count_occ(inp)
    job = execute(qc,Aer.get_backend('qasm_simulator'),shots=1000)
    counts = job.result().get_counts(qc) 
    assert counts[result]==1000
### END HIDDEN TESTS

```