<a href="https://colab.research.google.com/github/sbbzplt/QiskitFallFest_IZTECH/blob/master/Question_1_Final_Q.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Instructions

The places where you have enter code, to answer the questions, are marked with `# YOUR CODE HERE`. Once you have written your code you should remove the `raise NotImplementedError()` statement.

## Question 1

Suppose that you are given a quantum circuit with 8 qubits. The first 6 qubits are initialized so that each pair of qubits is representing a number in binary. 

i.e. qubits 0,1 represent the first number, qubits 2,3 represent the second number, qubits 4,5 represent the third number.

Given an integer $0\leq k \leq 3$ represented in binary format (i.e. using 2 bits), your task is to check whether $k$ exists within the 3 numbers you are given.

Example:

If the first 6 qubits are storing 110111 and $k$=11.
#1 = 11
#2 = 01
#3 = 11

So 11 exists among the three numbers.

Implement the function `check_k` which stores 1 in qubit numbered 6 if $k$ exists, and 0 otherwise.

You should try to use
- as few qubits as possible

Function `check_k` should have the following properties

* Inputs: A quantum circuit with 10 qubits, integer $k$ given as a binary string
* Returns: Quantum circuit

<!-- grade_id: cell-0f4849fc0faf37f8, points: 0, grade: false, solution: true -->
```python
from qiskit import *

def check_k(qc, k):
    ### BEGIN SOLUTION

    ### END SOLUTION
    qc.measure(6,0)
   
    return qc
    
```

```python
# You can use this code to test your function

qc = QuantumCircuit(10,1)

#initialize to 000111
qc.x(3)
qc.x(4)
qc.x(5)

#Let's check occurence of "11", you should see 1 measured for 1000 times 
qc = check_k(qc,"11")
job = execute(qc,Aer.get_backend('qasm_simulator'),shots=1000)
counts = job.result().get_counts(qc)   
print("Occurence of 11:", counts)

#Let's check occurence of "10", you should see 0 measured for 1000 times 
qc = check_k(qc,"10")
job = execute(qc,Aer.get_backend('qasm_simulator'),shots=1000)
counts = job.result().get_counts(qc)   
print("Occurence of 10:", counts)
```

Output:
```
Occurence of 11: {'1': 1000}
Occurence of 10: {'0': 1000}
```

```python
#You can use this cell to count the number of operations
tp = transpile(qc, basis_gates = ['cx', 'u'])
print(tp.count_ops())
```

Output:
```
OrderedDict([('u', 80), ('cx', 64), ('barrier', 2), ('measure', 2)])
```

<!-- grade_id: cell-8e7125cbfccba3b9, points: 0, grade: true, solution: false -->
```python
# hidden tests will be used for grading.
# If this cell results in an error, your implementation is incorrect
### BEGIN HIDDEN TESTS
def initialize(qc,input):
    for i in range(len(input)):
        if input[i]=="1":
            qc.x(i)
    return qc
        
def generate_binary_strings(bit_count):
    binary_strings = []
    def genbin(n, bs=''):
        if len(bs) == n:
            binary_strings.append(bs)
        else:
            genbin(n, bs + '0')
            genbin(n, bs + '1')
    genbin(bit_count)
    return binary_strings

def check_occ(inp,k):
    for i in range(3):
        if inp[2*i]==k[0] and inp[2*i+1] == k[1]:
            return "1"
    return "0"

inputs = generate_binary_strings(6)
k_list = ["00","01","10","11"] 
for inp in inputs:
    for k in k_list: 
        qc = QuantumCircuit(10,1)
        initialize(qc,inp)
        qc = check_k(qc,k)
        result = check_occ(inp,k)
        job = execute(qc,Aer.get_backend('qasm_simulator'),shots=1000)
        counts = job.result().get_counts(qc) 
        assert counts[result]==1000

tp = transpile(qc, basis_gates = ['cx', 'u'])
print(tp.count_ops())
### END HIDDEN TESTS

```

Output:
```
OrderedDict([('u', 43), ('cx', 32), ('barrier', 1), ('measure', 1)])
```